```python
# Initialize Otter
import otter
grader = otter.Notebook("test1.ipynb")
```

<a href="https://colab.research.google.com/github/hernansalinas/autogradesMetodosComputacionales/blob/main/Actividades_clase/actividad03/test1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
import otter
import pandas as pd
import numpy as np
%matplotlib inline
import otter
grader = otter.Notebook()
```

Output:
```
File /usr/local/lib/python3.10/dist-packages/IPython/core/interactiveshell.py:3658, in InteractiveShell.enable_matplotlib(self, gui)
   3637 """Enable interactive matplotlib and inline figure support.
   3638 
   3639 This takes the following steps:
   (...)
   3655     display figures inline.
   3656 """
   3657 from IPython.core import pylabtools as pt
-> 3658 gui, backend = pt.find_gui_and_backend(gui, self.pylab_gui_select)
   3660 if gui != 'inline':
   3661     # If we have our first gui selection, store it
   3662     if self.pylab_gui_select is None:

File /usr/local/lib/python3.10/dist-packages/IPython/core/pylabtools.py:320, in find_gui_and_backend(gui, gui_select)
    303 def find_gui_and_backend(gui=None, gui_select=None):
    304     """Given a gui string return the gui and mpl backend.
    305 
    306     Parameters
   (...)
    317     'WXAgg','Qt4Agg','module://matplotlib_inline.backend_inline','agg').
    318     """
--> 320     import matplotlib
    322     has_unifi
```

**Question 1:** Escribir una función llamada `promedio` que retorne el promedio de una arrray de numpy 

```python
def promedio(x):
    ...
```

```python
grader.check("q1")
```

**Question 2:** Create a DataFrame mirroring the table below and assign this to `data`. Then group by the `flavor` column and find the mean price for each flavor; assign this **series** to `price_by_flavor`.

| flavor | scoops | price |
|-----|-----|-----|
| chocolate | 1 | 2 |
| vanilla | 1 | 1.5 |
| chocolate | 2 | 3 |
| strawberry | 1 | 2 |
| strawberry | 3 | 4 |
| vanilla | 2 | 2 |
| mint | 1 | 4 |
| mint | 2 | 5 |
| chocolate | 3 | 5 |

```python
df=pd.read_excel("https://docs.google.com/spreadsheets/d/e/2PACX-1vSdJD5q-o-D3CriPHfkio4JJu0l4BJG4IIDTjREi2ytCUnT2CPLh5061vz1tAUZ92XQDOc8m_doks98/pub?output=xlsx")
df=df.set_index(df.flavor)
df
```

```python
data = ...
price_by_flavor = ...
price_by_flavor
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

**Question 3:** Create a barplot of `price_by_flavor`.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 4:** What do you notice about the bar plot?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> x_ = np.arange(1, 100)
>>> promedio(x_) == x_.mean()
True
```


Tests for q2 (points: None)
```
>>> len(data['flavor'].unique()) == 4
True
```
```
>>> for l in ['chocolate', 'vanilla', 'strawberry', 'mint']:
...     assert l in data['flavor'].unique()
```
```
>>> assert type(price_by_flavor) == pd.Series
```
```
>>> assert len(price_by_flavor) == 4
```
