```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab04_pandas.ipynb")
```

```python
import otter
import pandas as pd
import numpy as np
%matplotlib inline
import otter
import matplotlib.pylab as plt
grader = otter.Notebook()

```


**Pregunta  1:** Para los planetas de nuestro sistema solar, construir una serie de pandas  `ds` con indice nombre del planeta y valor la distancia al sol en UA.

Puedes intentanar generarlos con una IA, ejemplo de prompt:

hint:
`Generar un diccionario cuya clave es el nombre del planeta y valor la distancia al sol en UA`


a. Acceder al valor de mercurio y almacenarlo en una variable llamada dM.

b. Acceder al valor de Jupyter y almacenarlo en una variable llamada dJ.

c. Generar una nueva serie con los planetas cuya distancia es menor a 10UA, almacenarla en s

```python

ds = ...
dM = ...
dJ = ...
s = ...
ds
```

Output:
```
Ellipsis
```

```python
grader.check("q1")
```

Output:
```
q1 results:
    q1 - 1 result:
        ❌ Test case failed
        Trying:
            isinstance(ds, pd.Series) == True
        Expecting:
            True
        **********************************************************************
        Line 1, in q1 0
        Failed example:
            isinstance(ds, pd.Series) == True
        Expected:
            True
        Got:
            False

    q1 - 2 result:
        ❌ Test case failed
        Trying:
            np.isclose(dM, 0.39, rtol=0.001) == True
        Expecting:
            True
        **********************************************************************
        Line 1, in q1 1
        Failed example:
            np.isclose(dM, 0.39, rtol=0.001) == True
        Exception raised:
            Traceback (most recent call last):
              File "/Users/neuman/anaconda3/lib/python3.10/doctest.py", line 1350, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q1 1[0]>", lin
```

**Pregunta 2:**

a. Construir una serie de pandas  `dt`, con indice nombre del planeta y valor el periodo orbital en unidades de periodos terrestres.

b. Construir un dataframe de  pandas `df` con la serie ds y la serie dt, los nombres de las columnas deberán ser:
 `DistanciaSolUA`, `PeriodoUT`, el indice debe ser el nombre del planeta


```python

dt = ...
df = ...
df
```

```python
grader.check("q2")
```

**Pregunta 3:** El dataset [Temperature_Bogota_Barcelona](https://github.com/hernansalinas/Curso_Metodos_Computacionales/blob/main/Entregas/Parciales/Parcial1/dataset/Temperature_Bogota_Barcelona.csv)
 contiene información de la temperatura (en grados Fahrenheit) promedio de Bogotá y Barcelona desde el año 1995. Para el presente dataset realizar lo siguiente:


 a. Leer el dataset desde el link sugerido.

 b. Eliminar las columnas Month,  day y Unnamed

 c. Adicionar una nueva columna llamada `AvgTempCelcius` con la temperatura en grados celcius. $ C=(F-32)/1.8$



```python
df = ...
...
...

df
```

```python
grader.check("q3")
```

**Pregunta 4** Para el miso data set anterior realizar lo siguiente:

 a. Crear un nuevo data frame llamado df1, con los valores de temperatura de bogota que sean menores que 20 °C.
 b. Crear un nuevo data frame llamado df2, con los valores de temperatura de Barcelona que este entre 10 °C y 20 °C

```python
df1 = ...
...
df2 = ...
...
print(df1.head())
print(df2.head())
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

**Pregunta  5** Para el data set df, con todos los valores de Bogota y Barcelona, realizar en misma figura lo siguiente, recuerda poner las unidades y etiquetas respectivas :

 d. Gráfico de AvgTempCelcius como función del año para la ciudad de Bogotá.

 e. Gráfico AvgTempCelcius como funcion del año para la ciudad de Barcelona.

```python

plt.figure()
...
...
...
...
...
...

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Pregunta 5:** ¿Qué puede inferir del gráfico?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> isinstance(ds, pd.Series) == True
True
```
```
>>> np.isclose(dM, 0.39, rtol=0.001) == True
True
```
```
>>> np.isclose(dJ, 5.2, rtol=0.001) == True
True
```
```
>>> (len(ds) == 9) == True
True
```
```
>>> (len(s) == 6) == True
True
```


Tests for q2 (points: None)
```
>>> isinstance(dt, pd.Series) == True
True
```
```
>>> (df.columns == ['DistanciaSolUA', 'PerdiodoUT']).sum() == 2
True
```
```
>>> (df.index == ['Mercurio', 'Venus', 'Tierra', 'Marte', 'Jupiter', 'Saturno', 'Urano', 'Neptuno', 'Pluton']).sum() == 9
True
```


Tests for q3 (points: None)
```
>>> (df.columns.sort_values() == ['AvgTempCelcius', 'AvgTemperature', 'City', 'Year']).sum() == 4
True
```
```
>>> np.isclose(df['AvgTempCelcius'].mean(), 14.566, rtol=0.001) == True
True
```


Tests for q4 (points: None)
```
>>> np.isclose(df1.AvgTempCelcius.mean(), 12.9515, rtol=0.0001) == True
True
```
```
>>> np.isclose(df2.AvgTempCelcius.mean(), 16.1818, rtol=0.0001) == True
True
```
